---
title       : Diskreetsed jaotused
description : Insert the chapter description here



--- type:NormalExercise lang:r xp:100 skills:1 key:4f1b6aac58
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
